---
title: "End Term Examination"
output: 
  learnr::tutorial:
     progressive: true
     allow_skip: true
highlight: textmate
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(data.table)
library(tidyverse)
library(formattable)
library(lubridate)
library(gradethis)
library(learnrhash)
library(sortable)

tutorial_options(exercise.reveal_solution = FALSE)

gradethis::gradethis_setup(
  grading_problem.message = "Thank you for the submission",
  pass = " Thank you for the submission ",
  fail = "Thank you for the submission",
  pipe_warning = NULL,
  fail.encourage = NULL
)

orders <- fread("order.csv")
customer <- fread("customer.csv")
inventory <- fread( "FG_inventory.csv")
supplier <- fread("Supplier_data.csv")
supply_risk <- fread("supply_risk.csv")



available_inventory <- inventory%>%
  summarise(inventory = sum(inventory))

customer_total <- orders%>%
  group_by(customer)%>%
  summarise(qty = sum(qty))


# Equal_share calculation
  
  equal_share <- customer %>%
      select(customer) %>%
      mutate(
        split_equal = 1 / n(),
        inventory = available_inventory$inventory,
        allocated_inventory = round(inventory * split_equal,0),
        type = "Equal Share"
      ) %>%
      select(type, customer, allocated_inventory)
  
  
    # Fair_share calculation

  fair_share <- orders %>%
    group_by(customer)%>%
    summarise(qty = sum(qty))%>%
      mutate(
        split_fair = qty / sum(qty),
        inventory = available_inventory$inventory,
        allocated_inventory = round(inventory * split_fair,0),
        type = "Fair Share"
      ) %>%
      select(type, customer, allocated_inventory)
  

    # Penalty prioritization
  
  penalty_priority <- customer %>%
      mutate(
        rank = rank(late_penalty),
        inventory = if_else(rank == 5, available_inventory$inventory, as.integer(0))
      ) %>%
      arrange(desc(rank)) %>%
      left_join(select(customer_total, customer, qty)) %>%
      mutate(
        EOH = as.integer(cumsum(inventory - qty)),
        allocated_inventory = round(if_else(EOH > 0, qty,
          if_else(abs(EOH) < qty, qty + EOH, as.integer(0))),
          0),
        type = "Penalty priority"
      ) %>%
      select(type, customer, allocated_inventory)
  
  
    # Margin Prioritization

  margin_priority <- customer %>%
      mutate(
        margin = sell_price - cogs,
        rank = rank(margin),
        inventory = if_else(rank == 5, available_inventory$inventory, as.integer(0))
      ) %>%
      arrange(desc(rank)) %>%
      left_join(select(customer_total, customer, qty)) %>%
      mutate(
        EOH = as.integer(cumsum(inventory - qty)),
        allocated_inventory = round(if_else(EOH > 0, qty,
          if_else(abs(EOH) < qty, qty + EOH, as.integer(0))
        ),0),
        type = "Margin priority"
      ) %>%
      select(type, customer, allocated_inventory)
  
  
   allocation_summary <- equal_share %>%
      bind_rows(fair_share) %>%
      bind_rows(penalty_priority) %>%
      bind_rows(margin_priority)

   final_allocation <- margin_priority%>%
  select(-type)
   
# for question 26
   
   supply_startweek <- supplier%>%
      mutate(
        delivery_startdate = mdy(delivery_startdate),
        week = isoweek(delivery_startdate)
      ) %>%
      select(-delivery_startdate)
   
   input_current_supplier <- 0.25
input_supplier1 <- 0.25
input_supplier2 <- 0.25
input_supplier3 <- 0.25
  
    
  week_no <- seq(from = 33, to = 37, by = 1)

    weeklysupply_function <- function(x) {
      supplier %>%
        select(supplier) %>%
        mutate(week = as.integer(x))
    }
    
    supply_allocation <-  purrr::map_dfr(week_no, weeklysupply_function)%>%
      left_join(select(supply_startweek, supplier, week, weekly_supply)) %>%
        arrange(supplier, week) %>%
        group_by(supplier) %>%
        fill(weekly_supply, .direction = "down") %>%
        mutate(allocated_supply = if_else(supplier == "current_supplier",
          as.integer(weekly_supply * input_current_supplier),
          if_else(supplier == "Supplier_1", as.integer(weekly_supply * input_supplier1),
            if_else(supplier == "Supplier_2", as.integer(weekly_supply * input_supplier2),
              as.integer(weekly_supply * input_supplier3)
            )
          )
        )) %>%
        replace(is.na(.), 0)
   

knitr::opts_chunk$set(echo = TRUE)
```

## Exam Notes

The exam is based on the scenario modeling, resiliency assessment and supplier risk assignment we have done in the class. 

1. For some questions you are expected to `fill the missing functions or code` to answer.

2. For some questions you are expected to analyze the data and choose the correct answer 

3. For some questions you are expected to sort the code to generate the given output or analyze the data to sort the correct answer 

4. For some questions complete the code to generate the given output.

5. Once you are completed ,generate the hash code in the `submit`section and copy the code in the Google form provided. It is important to copy the complete #hash code for answer correction. advised to use the buttons to select instead of manual selection.

The data `customer` `orders` `inventory` and `supplier` been loaded for you. 

```{r eval=FALSE}
# Following libraries are loaded for you

library(data.table)
library(tidyverse)
library(formattable)
library(lubridate)

# Following data stored in the mentioned variables and available for you

orders <- fread("order.csv")
inventory <- fread("FG_inventory.csv")
customer <- fread("customer.csv")
supplier <- fread("Supplier_data.csv")
supply_risk <- fread("supply_risk.csv")

```


## Part A - 1 Mark

Please complete the following exercises by filling appropriate functions

### Question 1

select the customer column from the `orders` data

```{r q1, exercise = TRUE}

orders%>%
  ---(customer)

```

```{r q1-solution}

orders%>%select(customer)

```


```{r q1-check}

grade_this()
  
```



### Question 2

from the `orders`data filter the DC_5 from the warehouse column

```{r q2, exercise = TRUE}

orders%>%
  ---

```


```{r q2-solution}

orders%>%filter(warehouse == "DC_5")

```


```{r q2-check}

grade_this()
  
```

### Question 3

Join the inventory data from `inventory` table to `orders` data

```{r q3, exercise = TRUE}

orders%>%
---(inventory, by = c("customer", "warehouse"))

```

```{r q3-solution}

orders%>%
left_join(inventory, by = c("customer", "warehouse"))

```


```{r q3-check}

grade_this()
  
```

### Question 4

Add a column called `margin` to the customer data

```{r q4, exercise = TRUE}

customer%>%
---(margin = --- - cogs)

```


```{r q4-solution}

customer%>%
mutate(margin = sell_price - cogs)

```


```{r q4-check}

grade_this()
  
```



### Question 5

convert the `customer`data from a wider format (4 column) to longer format (2 column)

```{r q5, exercise = TRUE}

customer%>%
---(!customer, names_to = "key", values_to = "value")

```

```{r q5-solution}

customer%>%
pivot_longer(!customer, names_to = "key", values_to = "value")

```


```{r q5-check}

grade_this()
  
```


### Question 6

convert the `supplier_modified` data from a longer format( 2 column ) to wider format (4 column)

```{r q6-data}


supplier_modified <- supplier%>%
  select(!delivery_startdate)%>%
  pivot_longer(!supplier, names_to = "key", values_to = "value")

```


```{r q6, exercise = TRUE, exercise.setup = "q6-data"}

supplier_modified%>%
  ---(names_from = "key",values_from = "value")

```

```{r q6-solution}

supplier_modified%>%
  pivot_wider(names_from = "key",values_from = "value")

```


```{r q6-check}

grade_this()
  
```


### Question 7

Note : You have only one chance to select the answer. once selected can't change it. 


![](images/image_q7.png)

```{r q7, echo=FALSE}

question("what is the type of input highlighted in the above UI ?",
         
         type = "single",
  answer("numericInput"),
  answer("dataInput"),
  answer("fileInput", correct = TRUE),
  answer("uploadInput"),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```





### Question 8

Note : You have only one chance to select the answer. once selected can't change it.

![](images/image_q8.png)

```{r q8,echo=FALSE}

question("what is the type of input highlighted in the above UI ?",
  answer("fileInput"),
  answer("numberInput"),
  answer("dataInput"),
  answer("numericInput", correct =  TRUE),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```



### Question 9 

Note : You have only one chance to select the answer. once selected can't change it.

![](images/image_q9.png)


```{r q9,echo=FALSE}

question("what is the type of output highlighted in the above UI ?",
  answer("boxOutput"),
  answer("valueOutput"),
  answer("valueboxOutput", correct =  TRUE),
  answer("impactOutput"),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```



### Question 10

Note : You have only one chance to select the answer. once selected can't change it.

CTL - customer tolerance lead time , PLT - Product lead time

```{r q10,echo=FALSE}

question("Under what condition the need to plan the demand and supply is must ?",
  answer("CTL = PLT"),
  answer("CTL < PLT"),
  answer("CTL > PLT", correct =  TRUE),
  answer("CTL & PLT > 6 months"),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```





## Part B - 2 Marks

### Question 11

From the `orders` data aggregate the order qty at customer level as below by filling the missing details. 


```{r, echo = FALSE, message=FALSE}

orders%>%
group_by(customer)%>%
  summarise(qty= sum(qty))

```



```{r q11, exercise = TRUE}

orders%>%
---(customer)%>%
  summarise(qty = ---(qty))

```


### Question 12

From the `inventory`data prepare the following summary by filling the missing functions

```{r, echo = FALSE, message=FALSE}

inventory%>%
group_by(warehouse)%>%
  summarise(inventory= sum(inventory))

```



```{r q12, exercise = TRUE}

inventory%>%
---(warehouse)%>%
  summarise(inventory = ---(---))

```

### Question 13

```{r, echo = FALSE, message=FALSE}

orders%>%
group_by(customer,warehouse)%>%
  summarise(qty = sum(qty))

```

Note : You have only one chance to select the answer. once selected can't change it.

```{r q13, echo = FALSE}

flow <- c("orders%>%",
"group_by(customer,warehouse)%>%",
  "summarise(qty = sum(qty))")

question_rank(
  "Sort the code correctly so that you get the above output",
  answer(flow, correct = TRUE),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```


### Question 14 

```{r, echo = FALSE, message=FALSE}

orders%>%
      mutate(
        req_date = mdy(req_date),
        week = isoweek(req_date)
      )

```

Note : You have only one chance to select the answer. once selected can't change it.

```{r q14, echo = FALSE}

flow <- c("orders%>%",
"mutate(",
  "req_date = mdy(req_date)",
"week = isoweek(req_date)",
")" )

question_rank(
  "Sort the code correctly so that you get the above output",
  answer(flow, correct = TRUE),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```


### Question 15 

```{r echo=FALSE, message=FALSE}

supplier%>%
      select(supplier, cost,co2_per_shipment)%>%
  filter(supplier == "Supplier_1")

```

Note : You have only one chance to select the answer. once selected can't change it.

```{r q15, echo = FALSE}

flow <- c("supplier%>%",
"select(filter, cost,co2_per_shipment)%>%",
  "filter(supplier == `Supplier_1`)"
)

question_rank(
  "Sort the code correctly so that you get the above output",
  answer(flow, correct = TRUE),
  allow_retry = TRUE,
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```

### Question 16

Note : You have only one chance to select the answer. once selected can't change it.

![](images/image_q16.png)

```{r q16,echo=FALSE}

question("what are the outputs available in the UI ?",
  answer(" valueboxOutput & tableOutput"),
  answer("formattableOutput & valueboxOutput",correct =  TRUE),
  answer("boxOutput & formattableOutput"),
  answer("tableOutput & boxOutput"),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```


### Question 17

Note : You have only one chance to select the answer. once selected can't change it.

![](images/image_q17.png)

```{r q17,echo=FALSE}

question("what are the outputs available in the UI ?",
  answer("tableOutput & plotOutput"),
  answer("formattableOutput & chartOutput"),
  answer("plotOutput & formattableOutput",correct =  TRUE),
  answer("tableOutput & chartOutput"),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```


### Question 18

Note : You have only one chance to select the answer. once selected can't change it.

![](images/image_q18.JPG)


```{r q18,echo=FALSE}

question("Which company do you think have better resiliency Impact ?",
  answer("Premium foods on Produce Breadth & Perfect Pasta on Supply Lead time"),
  answer("Premium foods on supply Lead time & Perfect Pasta on Produce Breadth"),
  answer("Premium foods on both the factors",correct =  TRUE),
  answer("Perfect pasta on both the factors"),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```


### Question 19

Note : You have only one chance to select the answer. once selected can't change it.

![](images/image_q19.JPG)


```{r q19,echo=FALSE}

question("Which company do you think have better resiliency Impact ?",
  answer("Premium foods on Sources of supply & Perfect Pasta on planning process"),
  answer("Premium foods on planning process & Perfect Pasta on Sources of supply",correct =  TRUE),
  answer("Premium foods on both the factors"),
  answer("Perfect pasta on both the factors"),
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```


### Question 20


```{r , echo=FALSE, message= FALSE}

orders%>%
  group_by(customer, warehouse)%>%
  summarise(qty = sum(qty))%>%
  left_join(select(inventory, customer, warehouse, inventory))%>%
  mutate(coverage = inventory/qty,
         stock_out = 1-coverage)



```


Generate the above table by filling the missing code

```{r q20,exercise = TRUE}

orders%>%
  group_by(customer, warehouse)%>%
  summarise(qty = sum(qty))%>%
  left_join(select(inventory, customer, warehouse, inventory))%>%
  ---(--- = ---/---,
         --- = ---)



```



## Part C - 5 Marks

### Question 21

Based on  the `orders` `inventory` and `customer` data. Calculate the revenue and penalty impact at each customer- warehouse level as below. 

```{r , echo=FALSE, message= FALSE}

orders %>%
  group_by(customer,warehouse) %>%
  summarise(qty = sum(qty)) %>%
  left_join(inventory, by=c("customer","warehouse"))%>%
  left_join(customer, by = "customer")%>%
   mutate(
    revenue_impact = (qty - inventory) * sell_price,
    penalty_impact = (qty - inventory) * late_penalty)

```


```{r Q21, exercise = TRUE}

orders %>%
  

```


### Question 22

Total inventory is stored in `available_inventory` for you. Please check in the console. Split the total inventory using fair share approach as below.

```{r , echo=FALSE, message= FALSE}
orders %>%
  group_by(customer)%>%
  summarise(qty = sum(qty))%>%
      mutate(
        split_fair = qty / sum(qty),
        inventory = available_inventory$inventory,
        allocated_inventory = round(inventory * split_fair,0),
        type = "Fair Share"
      ) 

```



```{r q22, exercise = TRUE}

orders %>%


```


### Question 23 

Fill the missing 8 parts of the code to generate the below table creating the margin based priority to allocate existing inventory. also `customer_total`is created for you and contains the total requested qty from each customer. 


```{r , echo=FALSE, message= FALSE}
customer %>%
      mutate(
        margin = sell_price - cogs,
        rank = rank(margin),
        inventory = if_else(rank == 5, available_inventory$inventory, as.integer(0))
      ) %>%
      arrange(desc(rank)) %>%
      select(customer, margin, rank, inventory)%>%
      left_join(customer_total) %>%
      mutate(
        EOH = as.integer(cumsum(inventory - qty)),
        allocated_inventory = round(if_else(EOH > 0, qty,
          if_else(abs(EOH) < qty, qty + EOH, as.integer(0))
        ),0),
        type = "Margin priority"
      )

```



```{r q23, exercise = TRUE}

customer %>%
      mutate(
        --- = sell_price - ---,
        --- = rank(---),
        --- = if_else(rank == 5, available_inventory$inventory, as.integer(0))
      ) %>%
      arrange(desc(---)) %>%
      select(customer, margin, rank, inventory)%>%
      ---(customer_total) %>%
      mutate(
        --- = as.integer(cumsum(inventory - qty)),
        --- = round(if_else(EOH > 0, qty,
          if_else(abs(EOH) < qty, qty + EOH, as.integer(0))
        ),0),
        type = "Margin priority")


```


### Question 24 

`allocation_summary` and `customer_total` available for you. you need to calculate the business impact based on the 4 factors `ontime_revenue` `revenue_loss` `penalty` and `profit` as below by filling missing part of the code

```{r , echo=FALSE, message= FALSE}

allocation_summary %>%
      left_join(select(customer_total, customer, qty)) %>%
      left_join(customer) %>%
      mutate(
        ontime_revenue = round((allocated_inventory * sell_price) / 10^6, 2),
        revenue_loss = round(((qty - allocated_inventory) * sell_price) / 10^6, 2),
        penalty = round(((qty - allocated_inventory) * late_penalty) / 10^6, 2),
        profit = round((allocated_inventory * (sell_price - cogs)) / 10^6, 2)
      )

```


```{r q24, exercise = TRUE}

allocation_summary %>%
      left_join(---) %>%
      left_join(customer) %>%
      mutate(
        ontime_revenue = round((--- * ---) / 10^6, 2),
        revenue_loss = round(((--- - ---) * ---) / 10^6, 2),
        penalty = round(((--- - ---) * ---) / 10^6, 2),
        profit = round((--- * (--- - ---)) / 10^6, 2)
      )


```


### Question 25

```{r , echo=FALSE, message= FALSE}

orders %>%
      mutate(
        req_date = mdy(req_date),
        week = isoweek(req_date)) %>%
      group_by(customer, week) %>%
      summarise(qty = sum(qty)) %>%
      left_join(final_allocation) %>%
      mutate(
        allocated_inventory = if_else(week == min(week), allocated_inventory, 0),
        Net_requirement = cumsum(allocated_inventory - qty)) %>%
      rowwise() %>%
      mutate(unfulfilled_qty = if_else(Net_requirement < 0, min(abs(Net_requirement), qty), 0))

```


you can check the `final_allocation` data here

```{r fin_all, exercise = TRUE}

final_allocation
```



```{r q25, echo = FALSE}

flow <- c("orders %>%",
      "mutate(req_date = mdy(req_date),",
        "week = isoweek(req_date)) %>%",
      "group_by(customer, week) %>%",
      "summarise(qty = sum(qty)) %>%",
      "left_join(final_allocation) %>%",
      "mutate(allocated_inventory = if_else(week == min(week), allocated_inventory, 0),",
        "Net_requirement = cumsum(allocated_inventory - qty)) %>%",
      "rowwise() %>%",
      "mutate(unfulfilled_qty = if_else(Net_requirement < 0, min(abs(Net_requirement), qty), 0))"
)

question_rank(
  "Sort the code correctly so that you get the above output",
  answer(flow, correct = TRUE),
  allow_retry = TRUE,
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```


### Question 26


```{r , echo=FALSE, message= FALSE}

supply_allocation%>%
      group_by(supplier) %>%
      summarise(no_of_shipment = sum(allocated_supply != 0)) %>%
      left_join(supplier) %>%
      mutate(
        co2 = no_of_shipment * co2_per_shipment,
        plastics = no_of_shipment * plastics_per_shipment
      ) %>%
      summarise(
        co2 = sum(co2),
        plastics = sum(plastics)
      ) %>%
      mutate(type = "Post_selection") %>%
      select(type, co2, plastics)

```


you can check the `supply_allocation` data here

```{r sup_all, exercise = TRUE}

supply_allocation

```



```{r q26, echo = FALSE}

flow <- c("supply_allocation%>%",
      "group_by(supplier) %>%",
      "summarise(no_of_shipment = sum(allocated_supply != 0)) %>%",
      "left_join(supplier) %>%",
      "mutate(co2 = no_of_shipment * co2_per_shipment,",
        "plastics = no_of_shipment * plastics_per_shipment) %>%",
      "summarise(co2 = sum(co2),",
        "plastics = sum(plastics)) %>%",
      "mutate(type = `Post_selection`) %>%",
      "select(type, co2, plastics)"
)

question_rank(
  "Sort the code correctly so that you get the above output",
  answer(flow, correct = TRUE),
  
  allow_retry = TRUE,
  
  correct = "Thank you for the submission",
  incorrect = "Thank you for the submission"
)

```



## Part - D 15 Marks

The `supply_risk` data given to you which contains the list of components,  their suppliers and regions with respective spend value given. By analyzing the data answer the following questions.

```{r , echo=FALSE, message= FALSE}

head(supply_risk)

```

You can perform the necessary analysis here

```{r sup_risk, exercise = TRUE}


```

### Question 27 - 3 Marks

Identify the top two components have more supply sources(Supplier -region combination) ?

Note : Expected answer format "RMC1&RMC2". Stick to that. 

```{r q27, exercise = TRUE}


```


```{r q27-solution}

"RMC18&RMC20"

```


```{r q27-check}

grade_this()
  
```


### Question 28 - 3 Marks

During a supply disruption if the supplier is shut down , which of the following component is under more risk ?
RMC20/RMC15/RMC41/RMC14

Note : Expected answer format "RMC1". Stick to that.

```{r q28, exercise = TRUE}


```


```{r q28-solution}

"RMC14"

```


```{r q28-check}

grade_this()
  
```



### Question 29 - 3 Marks

During a supply disruption if NAM and APAC region affected and all the suppliers in these region unable to operate their plant. How many components in each region turns 100% risk of supply ?

Note : Expected answer format "1 & 2 respectively". Stick to that.


```{r q29, exercise = TRUE}


```


```{r q29-solution}

"3 & 6 respectively"

```


```{r q29-check}

grade_this()
  
```



### Question 30 - 3 Marks

In a given disruption if the existing source of supply is completely shut. How many components turns 100% risk of supply ?

Note : Expected answer format 1 (numeric). Stick to that


```{r q30, exercise = TRUE}


```


```{r q30-solution}

25

```


```{r q30-check}

grade_this()
  
```



### Question 31 - 3 Marks

In a given disruption LATAM region is completely affected. How many components supplied from LATAM supplier can utilize alternate source of supply ?

Note : Expected answer format 1 (numeric). Stick to that


```{r q31, exercise = TRUE}


```


```{r q31-solution}

6

```


```{r q31-check}

grade_this()
  
```





## Submit



```{r context="server"}
learnrhash::encoder_logic()
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui(ui_before = default_ui(url = "https://forms.gle/Xmqe5reyAUdGJnH18")
)
```


